## Simulating Monty Hall Problem

In this problem, we will be writing pieces to simulate the [Monty hall problem](https://en.wikipedia.org/wiki/Monty_Hall_problem). Monty Hall was the host of "Let's Make a Deal" in the 1960's and 70's.

The problem can be summarized by the following:

> Suppose you're on a game show, and you're given the choice of three doors: Behind one door is a car; behind the others, goats. You pick a door, say No. 1, and the host, who knows what's behind the doors, opens another door, say No. 3, which has a goat. He then says to you, "Do you want to pick door No. 2?" Is it to your advantage to switch your choice?

> 1. The host must always open a door that was not picked by the contestant.
1. The host must always open a door to reveal a goat and never the car.
1. The host must always offer the chance to switch between the originally chosen door and the remaining closed door.


The typical problem is for 3 doors but we can easily extend it to any number of doors. The simulation for this problem can be designed in different ways. However, we will focus our design using Python's dictionary and string objects along with some basic concepts of probability theory. 

**Caveat for the whole notebook** - This problem specifically deals with probabilities and simulations. Wherever required, we have implemented 99.7% confidence intervals for checking your answers. Do keep in mind that there is a probability (very small) that your answer lies beyond this interval. In such a scenario, you might want to re-run the test cell. 

**Exercise 0** (2 points) Complete the function `initialize_game()` which has 1 input argument:
- `n`: The total number of doors

The function will place a car behind one of the **n** doors **randomly and uniformly** (following Uniform Distribution).  

The function returns a dictionary with:

- keys as strings in the pattern 'Door _x_', where $1 \leq x \leq n$, and 
- values as either the string 'Goat' or the string 'Car'.

There can only be one 'Car' in the dictionary. All the other values shall be 'Goat'.

**Example**: For `n = 5`, `'Car'` can be behind any door from `'Door 1'` to `'Door 5'`. In this case, `'Car'` is behind `'Door 2'`, the function should return:   
`{'Door 1': 'Goat', 'Door 2': 'Car', 'Door 3': 'Goat', 'Door 4': 'Goat', 'Door 5': 'Goat'}` 

```python
def initialize_game(n):
    from random import randrange
    #
    # YOUR CODE HERE
    #

```

<!-- grade_id: test_ex_0, points: 2, grade: true, solution: false -->
```python
## Test code
import numpy as np
# Testing for 3 Doors
door_dict = initialize_game(3)
assert type(door_dict) is dict, "Make sure returned object is a dictionary"
assert len(door_dict) == 3, "Incorrect number of keys in dictionary"
assert 'Door 1' in door_dict and 'Door 2' in door_dict and 'Door 3' in door_dict, "Some keys missing in returned dictionary"

d1 = d2 = d3 = 0


for  i in range(10000):
    door_dict = initialize_game(3)
    if door_dict['Door 1'] == 'Car':
        d1 += 1/10000
    elif door_dict['Door 2'] == 'Car':
        d2 += 1/10000
    else:
        d3 += 1/10000

assert  0.36 > d1 and d1 >= 0.30, 'Function initialize_game() might not be doing correct random assignments'
assert  0.36 > d2 and d2 >= 0.30, 'Function initialize_game() might not be doing correct random assignments'
assert  0.36 > d3 and d3 >= 0.30, 'Function initialize_game() might not be doing correct random assignments'


# Testing for 5 Doors
door_dict = initialize_game(5)
assert type(door_dict) is dict, "Make sure returned object is a dictionary"
assert len(door_dict) == 5, "Incorrect number of keys in dictionary"
assert 'Door 1' in door_dict and 'Door 2' in door_dict and 'Door 3' in door_dict and 'Door 4' in door_dict and 'Door 5' in door_dict, "Some keys missing in returnd dictionary"

d = [0]*5
for  i in range(10000):
    door_dict = initialize_game(5)
    if door_dict['Door 1'] == 'Car':
        d[0] += 1/10000
    elif door_dict['Door 2'] == 'Car':
        d[1] += 1/10000
    elif door_dict['Door 3'] == 'Car':
        d[2] += 1/10000
    elif door_dict['Door 4'] == 'Car':
        d[3] += 1/10000
    else:
        d[4] += 1/10000

for i in range(len(d)):
    assert 0.24 > d[i] and d[i] >= 0.16, 'Function might not be intializing game following uniform distribution'

print("\nPassed!")
```

**Exercise 1** (1 point) Complete the function `first_choice()` which has 1 input argument: 

- `door_dict`:  A dictionary representing every numbered door and what that door is hiding (a dictionary similar to that returned by function `intitialize_game(n)`)

The goal here is to create a function where the player decides his/her first choice of door **randomly and uniformly** (following Uniform Distribution). 

The function returns `int` value indicating the door number selected by player depending on number of doors in `door_dict`

```python
def first_choice(door_dict):
#
# YOUR CODE HERE
#

```

<!-- grade_id: test_ex_1, points: 1, grade: true, solution: false -->
```python
# Test code
door_dict = {'Door 1': 'Car', 'Door 2': 'Goat', 'Door 3': 'Goat'}
door = first_choice(door_dict)
assert type(door) is int, 'Expected int found {}'.format(type(door))
# Test for 3 doors
d = [0]*3
for  i in range(10000):
    door = first_choice(door_dict)
    d[door - 1] += 1/10000
for i in range(3):
    assert  0.36 >= d[i] and d[i]>= 0.30, 'Function might not be choosing players first choice uniformly for 3 door game'

# Testing for 5 Doors
d = [0]*5
door_dict = {'Door 1': 'Car', 'Door 2': 'Goat', 'Door 3': 'Goat', 'Door 4': 'Goat', 'Door 5': 'Goat'}
for  i in range(10000):
    door = first_choice(door_dict)
    d[door -1] += 1/10000
for i in range(5):
    assert  0.25 > d[i] and d[i] >= 0.15, 'Function might not be choosing players first choice uniformly for 5 door game'

print("\nPassed!")
```

**Exercise 2** (3 points) Complete the function `monty_opens()` which has 3 input arguments. 

- `door_dict`: A dictionary similar to that returned by function `initialize_game()`
- `player_door`: An integer value denoting the player's first choice of door
- `k`: The number of doors to be opened by Monty

The function simulates Monty opening a door after the player has already picked an unopened door as their first choice. Opening a door is modelled by deleting goat doors from dictionary `door_dict`. 

The function returns the modified `door_dict`.

There are 2 situations you need to consider:

- **Case 1**: If the player's first choice has a `Goat` behind the door, then Monty will open **k** goat doors out of the remaining **n-2** goat doors randomly and uniformly. (Each goat door except player door equally likely to be opened). 

- **Case 2**: If the player's first choice has a `Car` behind the door, then Monty will open **k** goat doors out of the remaining **n-1** goat doors randomly and uniformly. (Each goat door equally likely to be opened). 

**Note:** Player door and car door can never be opened. 

**Example**: Given input:
- `door_dict` = `{'Door 1': 'Goat', 'Door 2': 'Car', 'Door 3': 'Goat', 'Door 4': 'Goat', 'Door 5': 'Goat'}`
- The player's first choice is `player_door` = 1 (meaning 'Door 1')
- `k = 2`, means Monty will open 2 doors out of `n-2=3` goat doors, which are Door 3, 4 and 5.

Therefore, your function should return  `door_dict = {'Door 1': 'Goat', 'Door 2': 'Car', 'Door x': 'Goat'}`.
- If Monty opens Door 3 and 4 -> x = 5    
- If Monty opens Door 3 and 5 -> x = 4    
- If Monty opens Door 4 and 5 -> x = 3   


```python
import random

def monty_opens(door_dict, player_door, k):
    assert type(door_dict) is dict, "input variable door_dict should be a dictionary"
    assert k <= len(door_dict) - 2, "Incorrect number of doors to be opened in dictionary"
#
# YOUR CODE HERE
#
    return door_dict
```

<!-- grade_id: test_ex_2, points: 3, grade: true, solution: false -->
```python
##### Test code

# Case 1: Player's first choice has Goat in 3 door Monty hall Game
door_dict = {'Door 1': 'Goat', 'Door 2': 'Car', 'Door 3': 'Goat'}
door_dict = monty_opens(door_dict, 1, 1)
assert type(door_dict) is dict, "Make sure returned object is a dictionary"
assert len(door_dict) == 2, "Incorrect number of keys in dictionary"
assert 'Door 1' in door_dict and 'Door 2' in door_dict, "Incorrect key removed"
door_dict = {'Door 1': 'Car', 'Door 2': 'Goat', 'Door 3': 'Goat'}
door_dict = monty_opens(door_dict, 3, 1)
assert len(door_dict) == 2, "Incorrect number of keys in dictionary"
assert 'Door 1' in door_dict and 'Door 3' in door_dict, "Incorrect key removed"

# Case 2: Player's first choice has Car
d1 = 0
d3 = 0
for i in range(10000):
    door_dict = {'Door 1': 'Goat', 'Door 2': 'Car', 'Door 3': 'Goat'}
    player_door = 2
    door_dict = monty_opens(door_dict, 2, 1)
    if 'Door 1' in door_dict:
        d1 += 1/10000
    else:
        d3 += 1/10000
assert  0.53 > d1 and d1 >= 0.47, 'Make sure that door is being opened randomly and uniformly'
assert  0.53 > d3 and d3 >= 0.47, 'Make sure that door is being opened randomly and uniformly'

# Case 3: Checking Random number of intial doors and random number of doors opened by Monty
import random
n_doors = random.randint(8, 13) 
k = random.randint(3, 6) 
car_door = 'Door ' + str(n_doors) 
player_door = 'Door 1'
d = [0]*n_doors
for i in range(10000):
    door_dict = {}
    for i in range(n_doors):
        door_dict["Door " + str(i + 1) ] = "Goat"
    door_dict["Door " + str(n_doors) ] = "Car"
    door_dict = monty_opens(door_dict, 1, k)
    for key, val in door_dict.items():
        d[int(key[5:]) - 1] += 1
assert d[0] == 10000, 'Player door should never be removed from door dict'
assert d[n_doors -1] == 10000, 'Car door should never be removed from doo dict'
for i in range(1, n_doors - 1):
    assert  k/(n_doors - 2)*(0.8) <= (1 -  d[i]/10000) <= k/(n_doors - 2)*(1.2), 'Doors not opening not uniformly and randomly'

print("\nPassed!")
```

**Exercise 3** (2 points) Complete the function `final_decision()` which has 3 input arguments:

- `door_dict`: A dictionary after Monty has opened `k` doors using `monty_opens()` function
- `player_door`: An integer value denoting player's first choice
- `decision`: A string value, either `'Stay'` or `'Switch'`, depending on whether player decides to stay with first choice or switch to another door

The function returns string `'Won'` if player wins the game and string `'Lost'` if he/she loses, depending on the player's decision to switch the door or not. 

The function simulates either: 

- staying with the first choice of doors, in which case the first choice door (`player_door`) is opened, and either a car or a goat is revealed, or
- switching to another door, in which case any of the doors that were NOT the first choice could be selected uniformly and randomly. The selected door is then opened, and either a car or a goat is revealed.

```python
def final_decision(door_dict, player_door, decision):
    #
    # YOUR CODE HERE
    #

```

<!-- grade_id: test_ex_3, points: 2, grade: true, solution: false -->
```python
##### Test code

# Case 1
door_dict = {'Door 1': 'Goat', 'Door 2': 'Car', 'Door 4': 'Goat'}
result = final_decision(door_dict, 1, 'Stay')
# General return statement check
assert result is 'Lost', "Please check the return value"

# Case 2
door_dict = {'Door 1': 'Goat', 'Door 2': 'Car', 'Door 4': 'Goat'}
result = final_decision(door_dict, 2, 'Stay')
assert result is 'Won', "Please check the return value"

# Case 3
door_dict = {'Door 1': 'Goat', 'Door 2': 'Car', 'Door 4': 'Goat','Door 5': 'Goat','Door 6': 'Goat'}
result = final_decision(door_dict, 2, 'Switch')
assert result is 'Lost', "Please check the return value"

# Case 4
door_dict = {'Door 1': 'Goat', 'Door 2': 'Car', 'Door 4': 'Goat','Door 5': 'Goat','Door 6': 'Goat'}
result_won = 0
result_lost = 0
for i in range(10000):
    result = final_decision(door_dict, 1, 'Switch')
    if result == "Lost":
        result_lost += 1
    else:
        result_won += 1
prob_winning = result_won/10000
#print(prob_winning)
assert prob_winning <= 0.27 and prob_winning >= 0.23 , "Doors not selected at random"

print("\nPassed!")
```

**Exercise 4** (2 points) Now, let's bring the pieces together and simulate a game. In this exercise, you will complete the given `simulate_game` function that simulates Monty Hall game 10000 times. The function has 2 input arguments:

- `n`: Number of doors
- `k`: Number of doors to be opened be Monty

The function returns a tuple with the number of wins each for staying and switching:

- `stay_won`: An integer count of the number of wins resulting from a `'Stay'` decision
- `switch_won`: An integer count of the number of wins resulting from a `'Switch'` decision

**Hint**: You have to think about how you will use the following functions - `first_choice` , `initialize_game` , `final_decision` and `monty_opens` and what their input parameters should be to complete the exercise.

```python
def simulate_game(n, k):
    stay_won = 0  #Variable to keep a track of the number of times the player won when he/she decided to stay
    switch_won = 0  #Variable to keep a track of the number of times the player won when he/she decided to switch
    
    for i in range(10000):
        #
        # YOUR CODE HERE
        #
    return(stay_won, switch_won)
```

<!-- grade_id: test_ex_4, points: 2, grade: true, solution: false -->
```python
##### Test code

#Case 1
stay_won, switch_won = simulate_game(3,1)
prob_switch_won = switch_won/10000

assert prob_switch_won <= 0.7  and prob_switch_won >= 0.64

#Case 2
stay_won, switch_won = simulate_game(10,2)
prob_switch_won = switch_won/10000

assert prob_switch_won <= 0.142  and prob_switch_won >= 0.10


#Case 3

n_doors = random.randint(8, 13) 
k = random.randint(3, 6) 

stay_won, switch_won = simulate_game(n_doors, k)
prob_switch_won = switch_won/10000

population_prob =  (1/n_doors)*((n_doors-1)/ (n_doors-1 - k))

assert prob_switch_won < population_prob * 1.2   and prob_switch_won >= population_prob *(0.80)

print("\nPassed!")
```

**Fin!** You've reached the end of this problem. Don't forget to restart the kernel and run the entire notebook from top-to-bottom to make sure you did everything correctly. If that is working, try submitting this problem. (Recall that you must submit and pass the autograder to get credit for your work!)