# Part 1 of 2: Processing an HTML file

One of the richest sources of information is [the Web](http://www.computerhistory.org/revolution/networking/19/314)! In this notebook, we ask you to use string processing and regular expressions to mine a web page, which is stored in HTML format.

**The data: Yelp! reviews.** The data you will work with is a snapshot of a recent search on the [Yelp! site](https://yelp.com) for the best fried chicken restaurants in Atlanta. That snapshot is hosted here: https://cse6040.gatech.edu/datasets/yelp-example

If you go ahead and open that site, you'll see that it contains a ranked list of places:

![Top 10 Fried Chicken Spots in ATL as of September 12, 2017](https://cse6040.gatech.edu/datasets/yelp-example/ranked-list-snapshot.png)

**Your task.** In this part of this assignment, we'd like you to write some code to extract this list.

## Getting the data

First things first: you need an HTML file. The following Python code will download a particular web page that we've prepared for this exercise and store it locally in a file.

> If the file exists, this command will not overwrite it. By not doing so, we can reduce accesses to the server that hosts the file. Also, if an error occurs during the download, this cell may report that the downloaded file is corrupt; in that case, you should try re-running the cell.

```python
import requests
import os
import hashlib

if os.path.exists('.voc'):
    data_url = 'https://cse6040.gatech.edu/datasets/yelp-example/yelp.htm'
else:
    data_url = 'https://github.com/cse6040/labs-fa17/raw/master/datasets/yelp.htm'

if not os.path.exists('yelp.htm'):
    print("Downloading: {} ...".format(data_url))
    r = requests.get(data_url)
    with open('yelp.htm', 'w', encoding=r.encoding) as f:
        f.write(r.text)

with open('yelp.htm', 'r', encoding='utf-8') as f:
    yelp_html = f.read().encode(encoding='utf-8')
    checksum = hashlib.md5(yelp_html).hexdigest()
    assert checksum == "4a74a0ee9cefee773e76a22a52d45a8e", "Downloaded file has incorrect checksum!"
    
print("'yelp.htm' is ready!")
```

Output:
```
'yelp.htm' is ready!
```

**Viewing the raw HTML in your web browser.** The file you just downloaded is the raw HTML version of the data described previously. Before moving on, you should go back to that site and use your web browser to view the HTML source for the web page. Do that now to get an idea of what is in that file.

> If you don't know how to view the page source in your browser, try the instructions on [this site](http://www.wikihow.com/View-Source-Code).

**Reading the HTML file into a Python string.** Let's also open the file in Python and read its contents into a string named, `yelp_html`.

```python
with open('yelp.htm', 'r', encoding='utf-8') as yelp_file:
    yelp_html = yelp_file.read()
    
# Print first few hundred characters of this string:
print("*** type(yelp_html) == {} ***".format(type(yelp_html)))
n = 1000
print("*** Contents (first {} characters) ***\n{} ...".format(n, yelp_html[:7500]))
```

Output:
```
*** type(yelp_html) == <class 'str'> ***
*** Contents (first 1000 characters) ***
<!DOCTYPE html>
<!-- saved from url=(0079)https://www.yelp.com/search?find_desc=fried+chicken&find_loc=Atlanta%2C+GA&ns=1 -->
<html xmlns:fb="http://www.facebook.com/2008/fbml" class="js gr__yelp_com" lang="en"><!--<![endif]--><head data-component-bound="true"><meta http-equiv="Content-Type" content="text/html; charset=UTF-8"><link type="text/css" rel="stylesheet" href="./Best Fried chicken in Atlanta, GA - Yelp_files/css"><style type="text/css">.gm-style .gm-style-cc span,.gm-style .gm-style-cc a,.gm-style .gm-style-mtc div{font-size:10px}
</style><style type="text/css">@media print {  .gm-style .gmnoprint, .gmnoprint {    display:none  }}@media screen {  .gm-style .gmnoscreen, .gmnoscreen {    display:none  }}</style><style type="text/css">.gm-style-pbc{transition:opacity ease-in-out;background-color:rgba(0,0,0,0.45);text-align:center}.gm-style-pbt{font-size:22px;color:white;font-family:Roboto,Arial,san
```

Oy, what a mess! It will be great to have some code read and process the information contained within this file.

## Exercise (5 points): Extracting the ranking

Write some Python code to create a variable named `rankings`, which is a list of dictionaries set up as follows:

* `rankings[i]` is a dictionary corresponding to the restaurant whose rank is `i+1`. For example, from the screenshot above, `rankings[0]` should be a dictionary with information about Gus's World Famous Fried Chicken.
* Each dictionary, `rankings[i]`, should have these keys:
    * `rankings[i]['name']`: The name of the restaurant, a string.
    * `rankings[i]['stars']`: The star rating, as a string, e.g., `'4.5'`, `'4.0'`
    * `rankings[i]['numrevs']`: The number of reviews, as an **integer.**
    * `rankings[i]['price']`: The price range, as dollar signs, e.g., `'$'`, `'$$'`, `'$$$'`, or `'$$$$'`.
    
Of course, since the current topic is regular expressions, you might try to apply them (possibly combined with other string manipulation methods) find the particular patterns that yield the desired information.

```python
import re
import numpy as np

from bs4 import BeautifulSoup
from pprint import pprint
```

<!-- grade_id: rankings, points: 0, grade: false, solution: true -->
```python
#
soup = BeautifulSoup(yelp_html, 'html.parser')
soup.title
#
```

Output:
```
<title>Best Fried chicken in Atlanta, GA - Yelp</title>
```

<!-- grade_id: rankings, points: 0, grade: false, solution: true -->
```python
#
display(soup.title.name)
display(soup.title.parent.name)
display(soup.p)
display(soup.a)
#

```

Output:
```
'title'
'head'
<p>
                First, try refreshing the page and clicking Current Location again. Make sure you click <b>Allow</b> or <b>Grant Permissions</b> if your browser asks for your location. If your browser doesn't ask you, try these steps:
            </p>
<a href="https://www.yelp.com/search?find_desc=fried+chicken&amp;find_loc=Atlanta%2C+GA&amp;ns=1#header_find_form" rel="nofollow">
                        Skip to Search Form
                    </a>
```

```python
# setup functions
def extract_business_names (read_html):
    # To find the business names on the page
    soup = BeautifulSoup(read_html, 'html.parser')
    
    raw_biznames = list(soup.find_all("a", attrs={"class":"biz-name js-analytics-click"}))
    
    # Convert the html format into strings
    all_biznames = [str(bizname) for bizname in raw_biznames]
    
    # Get indices for the start and end of the business names using regex
    beginnings = [re.search(r'><span>', bizname) for bizname in all_biznames]
    endings = [re.search(r'</span></a>', bizname) for bizname in all_biznames]
    # Get the end of the beginnings, and the beginning of the ends to find the indices where the bizname sits 
    biz_indices = list(zip([beg.span()[1] for beg in beginnings], [end.span()[0] for end in endings]))
    
    biznames = [name[biz_indices[i][0]: biz_indices[i][1]] for i, name in enumerate(all_biznames)]
    return biznames

def extract_avg_star_ratings (read_html):
    # Instantiate beautiful soup
    soup = BeautifulSoup(read_html, 'html.parser')
    
    # To find the stars for each of the listed business names (including ads)
    raw_stars = soup.find_all('div', attrs={'class':'i-stars'})
    str_stars = [str(star) for star in raw_stars]
    
    # Use regex to find the actual star ratings & isolate those
    beginnings = [re.search(r'title="', stars) for stars in str_stars]
    endings = [re.search(r' star rating', stars) for stars in str_stars]
    
    star_indices = list(zip([beg.span()[1] for beg in beginnings], [end.span()[0] for end in endings]))
    stars = [star[star_indices[i][0]: star_indices[i][1]] for i, star in enumerate(str_stars)]
    
    return stars

def extract_review_counts (read_html):
    # Instantiate BeautifulSoup
    soup = BeautifulSoup(read_html, 'html.parser')
    
    # find the review counts
    all_counts = soup.find_all('span', attrs={'class':'review-count'})
    str_counts = [str(count) for count in all_counts]
    
    # match the exact review count (only want the integer)
    digits = [re.search(r'\d+', str_counts[i]) for i in range(len(str_counts))]
    nums = [int(digit.group()) for digit in digits]
    
    return nums

def extract_price_ranges (read_html):
    soup = BeautifulSoup(read_html, 'html.parser')
    
    # Find all the price ranges
    all_pr_ranges = soup.find_all("span", attrs={'class':'business-attribute price-range'})
    str_pr_ranges = [str(prices) for prices in all_pr_ranges]
    
    # Find the beginning indices and the ending indices
    begs = [re.search(r'price-range">', ranges) for ranges in str_pr_ranges]
    ends = [re.search(r'</span>', ranges) for ranges in str_pr_ranges]
    # Zip the indices together
    indices = list(zip([beg.span()[1] for beg in begs], [end.span()[0] for end in ends]))
    
    # Get the price ranges from the string price ranges
    prices = [price[indices[i][0]: indices[i][1]] for i, price in enumerate(str_pr_ranges)]
    
    return prices

def get_only_results_not_ads (read_html):
    '''
    returns only the indices of the advertisements, so as to know which to exclude from the above functions
    '''
    soup = BeautifulSoup(read_html, 'html.parser')
    
    all_biznames = soup.find_all("a", attrs={"class":"biz-name js-analytics-click"})
    str_biznames = [str(bizname) for bizname in all_biznames]
    
    ad_indices = []
    for i, name in enumerate(str_biznames):
        if 'ad_business_id' in name:
            ad_indices.append(i) 
    
    return ad_indices

def remove_ads (idx_to_remove: list, result_attr:list) -> list:
    
    no_ads = []
    for i, result in enumerate(result_attr):
        if i not in idx_to_remove:
            no_ads.append(result)
        else:
            pass
        
    return no_ads
```

```python
biznames = extract_business_names(yelp_html)
star_ratings = extract_avg_star_ratings(yelp_html)
price_ranges = extract_price_ranges(yelp_html)
review_counts = extract_review_counts(yelp_html)

ad_indices = get_only_results_not_ads(yelp_html)

parent_d = dict(zip(['name', 'numrevs', 'price', 'stars'], [
    remove_ads(ad_indices, biznames),
    remove_ads(ad_indices, review_counts),
    remove_ads(ad_indices, price_ranges),
    remove_ads(ad_indices, star_ratings),
]))

pprint(parent_d)
```

Output:
```
{'name': ['Gus’s World Famous Fried Chicken',
          'South City Kitchen - Midtown',
          'Mary Mac’s Tea Room',
          'Busy Bee Cafe',
          'Richards’ Southern Fried',
          'Greens &amp; Gravy',
          'Colonnade Restaurant',
          'South City Kitchen Buckhead',
          'Poor Calvin’s',
          'Rock’s Chicken &amp; Fries'],
 'numrevs': [549, 1777, 2241, 481, 108, 93, 350, 248, 1558, 67],
 'price': ['$$', '$$', '$$', '$$', '$$', '$$', '$$', '$$', '$$', '$'],
 'stars': ['4.0',
           '4.5',
           '4.0',
           '4.0',
           '4.0',
           '3.5',
           '4.0',
           '4.5',
           '4.5',
           '4.0']}
```

```python
# create child dictionaries for each business result in the yelp search
# append each child dictionary to a list called `rankings`

rankings = []

for i in range(10):
    child_d = {}
    for k, v in parent_d.items():
        child_d[k] = v[i]
        
    rankings.append(child_d)
    
rankings
```

Output:
```
[{'name': 'Gus’s World Famous Fried Chicken',
  'numrevs': 549,
  'price': '$$',
  'stars': '4.0'},
 {'name': 'South City Kitchen - Midtown',
  'numrevs': 1777,
  'price': '$$',
  'stars': '4.5'},
 {'name': 'Mary Mac’s Tea Room',
  'numrevs': 2241,
  'price': '$$',
  'stars': '4.0'},
 {'name': 'Busy Bee Cafe', 'numrevs': 481, 'price': '$$', 'stars': '4.0'},
 {'name': 'Richards’ Southern Fried',
  'numrevs': 108,
  'price': '$$',
  'stars': '4.0'},
 {'name': 'Greens &amp; Gravy', 'numrevs': 93, 'price': '$$', 'stars': '3.5'},
 {'name': 'Colonnade Restaurant',
  'numrevs': 350,
  'price': '$$',
  'stars': '4.0'},
 {'name': 'South City Kitchen Buckhead',
  'numrevs': 248,
  'price': '$$',
  'stars': '4.5'},
 {'name': 'Poor Calvin’s', 'numrevs': 1558, 'price': '$$', 'stars': '4.5'},
 {'name': 'Rock’s Chicken &amp; Fries',
  'numrevs': 67,
  'price': '$',
  'stars': '4.0'}]
```

<!-- grade_id: rankings_test, points: 5, grade: true, solution: false -->
```python
# Test cell: `rankings_test`

assert type(rankings) is list, "`rankings` must be a list"
assert all([type(r) is dict for r in rankings]), "All `rankings[i]` must be dictionaries"

print("=== Rankings ===")
for i, r in enumerate(rankings):
    print("{}. {} ({}): {} stars based on {} reviews".format(i+1,
                                                             r['name'],
                                                             r['price'],
                                                             r['stars'],
                                                             r['numrevs']))

assert rankings[0] == {'numrevs': 549, 'name': 'Gus’s World Famous Fried Chicken', 'stars': '4.0', 'price': '$$'}
assert rankings[1] == {'numrevs': 1777, 'name': 'South City Kitchen - Midtown', 'stars': '4.5', 'price': '$$'}
assert rankings[2] == {'numrevs': 2241, 'name': 'Mary Mac’s Tea Room', 'stars': '4.0', 'price': '$$'}
assert rankings[3] == {'numrevs': 481, 'name': 'Busy Bee Cafe', 'stars': '4.0', 'price': '$$'}
assert rankings[4] == {'numrevs': 108, 'name': 'Richards’ Southern Fried', 'stars': '4.0', 'price': '$$'}
assert rankings[5] == {'numrevs': 93, 'name': 'Greens &amp; Gravy', 'stars': '3.5', 'price': '$$'}
assert rankings[6] == {'numrevs': 350, 'name': 'Colonnade Restaurant', 'stars': '4.0', 'price': '$$'}
assert rankings[7] == {'numrevs': 248, 'name': 'South City Kitchen Buckhead', 'stars': '4.5', 'price': '$$'}
assert rankings[8] == {'numrevs': 1558, 'name': 'Poor Calvin’s', 'stars': '4.5', 'price': '$$'}
assert rankings[9] == {'numrevs': 67, 'name': 'Rock’s Chicken &amp; Fries', 'stars': '4.0', 'price': '$'}

print("\n(Passed!)")
```

Output:
```
=== Rankings ===
1. Gus’s World Famous Fried Chicken ($$): 4.0 stars based on 549 reviews
2. South City Kitchen - Midtown ($$): 4.5 stars based on 1777 reviews
3. Mary Mac’s Tea Room ($$): 4.0 stars based on 2241 reviews
4. Busy Bee Cafe ($$): 4.0 stars based on 481 reviews
5. Richards’ Southern Fried ($$): 4.0 stars based on 108 reviews
6. Greens &amp; Gravy ($$): 3.5 stars based on 93 reviews
7. Colonnade Restaurant ($$): 4.0 stars based on 350 reviews
8. South City Kitchen Buckhead ($$): 4.5 stars based on 248 reviews
9. Poor Calvin’s ($$): 4.5 stars based on 1558 reviews
10. Rock’s Chicken &amp; Fries ($): 4.0 stars based on 67 reviews

(Passed!)
```

**Fin!** This cell marks the end of Part 1. Don't forget to save, restart and rerun all cells, and submit it. When you are done, proceed to Part 2.